# ITPPDS Mid Course Programming Project

### <span style="color:red">*Please read these instructions fully*</span>

### Marking Scheme

The four tasks in this notebook ask you to write functions to perform specified operations. Please complete the functions in the code cells provided, removing the `# YOUR CODE HERE` and `raise NotImplementedError()` text from the cell and replacing with your own code. Do not change the function names from those provided. These functions will be automatically marked by a computer, which will test whether your functions perform the required operations and award marks where those tests pass. You will receive partial marks if your code runs but fails to perform some of the requirements, but you will not receive any marks if you code produces an error or fails to return any valid values. It is critical that your functions are named as specified in the questions and that they return the values requested, without any additional output.

Please see [**this video**](https://media.ed.ac.uk/media/Completing+CoCalc+Assignments.mov/1_gt6elgym) for a full description of how to use the autograder cells to test your functions. We highly recommend using the `Validate` button to check your notebook for errors and to run the autograders prior to submission. The `Validate` button restarts the Python kernel and runs each cell in turn, so ensures that the cells run in the same order as they will when being graded.

The test is graded out of a total of 20 marks. The marks available for each question are details in individual mark schemes for each task.

### Plagiarism and collusion

This data analysis project forms a part of your course assessment and **must be your own work**. You are encouraged to use the self study notebooks, workshop notebooks, internet search and online Python documentation, and to use or adapt code examples from each of those sources.  However you must not help or receive help from anyone, and must not use code written by other students.

Where you have taken or adapted from other sources that is greater than a single line, please use Python comments (#) to give the source of any code unless it is from the course notebooks and workshop material.

All work will be checked for evidence of plagiarism and collusion. CoCalc notebooks retain a full file editing history and this may be reviewed.

### Deadline and submission information

This workbook will be automatically submitted at **2pm on Wednesday 2nd March**, which is the deadline for this assignment. You do not need to do anything to submit, other than ensure your notebook is saved (Click Save in the upper left hand side of this page). If you want to arrange for later collection, for example because you have been granted an extension to this assignment or would like to submit with late penalties applied, please email graeme.cowan@ed.ac.uk **before the deadline** so that alternative collection arrangements can be put in place.


<img style="float: right;" src="./Darwins_finches.png" width="500" height="350" >

# Geospiza Morphology

Geospiza are a genus of finches endemic to the Galapagos Islands that are also known as "Darwin's finches". A commonly used data science dataset is from experiments studying phenotypic variation among Geospiza ground-finches from six morphological measurements of museum specimens (wing length, tail length, tarsus length, bill length, bill width and bill depth) taken on adult males by H. S. Swarth for his monographic revision of the birds of the Galapagos (Swarth 1931). The combination of multiple measurements on each individual bird makes this dataset suitable for studying multivariate analyses and dimensionality reduction techniques. 

In the first part of this project, you are asked to write Python 3 functions that you might use to reformat or process phenotypic datasets such as the Geospiza morphology dataset.

# Task 1

Write a function called **unique_species()** that takes a list of species names as an argument and returns a list of all unique species sampled *(1 mark)* sorted in alphabetical order (*1 mark*). Your function should have a valid docstring *(1 mark)*. Note: the list of all unique species sampled must be sorted within your function.

A list of species is provided two cells below to enable you to test your function. The expected result from that list is `['G. fortis', 'G. magnirostris', 'G. propinqua']`, however your function must operate with any list of species names, not just the one provided.

##### Marking scheme (3 marks total)
* Function returns a list of unique species: 1 mark
* and is sorted alphabetically: 1 mark
* Function contains a valid docstring: 1 mark
* No marks will be awarded to this question if the function does not run, for example due to a syntax error in the cell.

You can check your answer with the Task 1 autograders 1-3 at the bottom of this assignment. If the message 'This test has passed' is displayed, you will receive the marks for that part of the question.


<!-- grade_id: Task1, points: 0, grade: false, solution: true -->
```python
# This cell must contain the function definition code only. Use the cell below for testing your function.
def unique_species(species):
    ### BEGIN SOLUTION
    '''takes a list of species as an argument and generates a list of unique species'''
    unique=[]
    for s in species:
        if not s in unique:
            unique.append(s)
    return sorted(unique)
    ### END SOLUTION

```

```python
# Example testing code for task 1. Please delete any additional code you add to this cell after testing.
species_list_a=['G. propinqua', 'G. fortis', 'G. propinqua', 'G. propinqua', 'G. propinqua', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis','G. magnirostris', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. magnirostris', 'G. fortis']

```

# Task 2

We have a list of Geospiza species and bill lengths reproduced from a subset of H.S. Swarth's measurement data. Write a function called **generate_dict_max()** that takes two lists as arguments: a list of species and a list of bill lengths (floats). Your function should return a dictionary with unique species names as keys and maximum bill lengths for each species as the associated values *(3 marks)*. Your function may optionally have a docstring (not required), but if present please use a single line docstring to facilitate autograding of this task.

You will gain 1 additional mark if your function, including the function definition line but excluding any docstring, contains six or fewer lines of code, or 2 marks if it contains only two lines. You have been taught enough Python to write the code in six or fewer lines - this simply requires an optimisation of the logic. To write the code in three or fewer lines will require a different approach to the problem and will likely need you to search web resources. n.b. Condensed code is not necessarily better code (often the opposite), however here this task is set to challenge your evaluation of different ways to achieve a task and to test your ability to use web resources. Please note that semi-colon line separators will be counted as newline characters.

Two lists are provided in the cell below the function definition to allow you to test your code. From the two lists provided, your results should be `{'G. propinqua': 22.5, 'G. magnirostris': 25.0, 'G. fortis': 19.5}`. The order of keys in the dictionary does not matter.

##### Marking scheme (5 marks total)
* Function generates a dictionary of maximum bill lengths (values) for each species (keys): 3 marks
* Function is 6 lines or fewer: 1 mark, or 2 lines or fewer: 2 marks (excluding the single line docstring)
* No marks will be awarded to this question if the function does not run, for example due to a syntax error in the cell.

You can check your answer with the Task 2 autograders 1-4 at the bottom of this assignment. If the message 'This test has passed' is displayed, you will receive the marks for that part of the question.


<!-- grade_id: Task2, points: 0, grade: false, solution: true -->
```python
# This cell must contain the function definition code only. Use the cell below for testing your function.
def generate_dict_max(species, lengths):
    ### BEGIN SOLUTION
    '''Generates dictionary of species as keys and maximum values as values'''
    # 1 line method
    #return dict(sorted(list(zip(species, lengths)), key=lambda x: x[1]) )

    # Another 1 line method
    return {s:v for (s,v) in sorted(list(zip(species, lengths)), key=lambda x: x[1])}

    # 5 line alternative
    #out_dict={}
    #for i, s in enumerate(species):
    #    if s not in out_dict or lengths[i] > out_dict[s]:
    #            out_dict[s]=values[i]
    #return out_dict

    ### END SOLUTION

```

```python
# Example testing code for Task 2. Please delete any additional code you add to this cell after your testing.
species_list=['G. propinqua', 'G. propinqua', 'G. propinqua', 'G. propinqua', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis']

bill_length_list=[19.5, 20.0, 22.5, 21.0, 25.0, 22.5, 21.5, 22.0, 24.0, 22.2, 21.0, 21.0, 22.5, 24.0, 22.0, 22.5, 24.0, 20.2, 20.0, 21.5, 17.0, 15.2, 16.0, 16.8, 16.0, 16.0, 16.5, 17.0, 16.5, 15.8, 15.8, 16.0, 18.5, 19.5, 15.8, 19.5, 17.0, 16.2, 16.0, 15.0, 15.0, 15.0, 16.5, 15.8, 13.0, 15.0, 15.0, 17.0, 15.2, 16.5]

```

# Managing Health Data - Community Health Index


The Community Health Index Number (CHI number) is an identifier used to uniquely identify patients within the NHS in Scotland ([Link to NHS Scotland description ](https://www.ndc.scot.nhs.uk/Data-Dictionary/SMR-Datasets//Patient-Identification-and-Demographic-Information/Community-Health-Index-Number/)). Data scientists at Public Health Scotland maintain a library of tools to manage health data written in the R language called [phsmethods](https://github.com/Public-Health-Scotland/phsmethods), and in this part of the project we will replicate some of those tools in Python. 

<img src="CHI_number.png" alt="CHI_number" style="float: left; margin-right: 25px;" />
The CHI number is composed of each of the following elements in order:
<br>
<br>

* Digits 1-6: Date of Birth in the format DDMMYY
* Digits 7&8: Two further digits
* Digit 9: an even number for female patients and an odd number for male patients
* Digit 10: A modulus 11 check digit

# Task 3

Write a Python 3 function named sex_from_chi() that takes a 10 digit CHI number as a **string** and returns the string 'Male' if the 9th digit in the CHI number is odd, and 'Female' if it is even (including zero) (2 marks). If too long or short a string is entered, your function should raise a ValueError, and if the CHI is entered as an integer value, the function should raise a TypeError. (2 marks). Your function must have a valid docstring (1 mark).

#### Mark scheme (5 marks total)

* Function returns 'Male' for odd 9th digit, 'Female' for even 9th digit - 2 marks
* Function returns ValueError if the string is of incorrect length and TypeError if the CHI is entered as an integer value - 2 marks
* Function has a valid docstring - 1 mark



<!-- grade_id: Task3, points: 0, grade: false, solution: true -->
```python
def sex_from_chi(chi): 
    ### BEGIN SOLUTION
    """
    Function takes a 10 digit CHI number as a **string** and returns the string 'Male' if the 9th digit in the CHI number is odd,     and 'Female' if it is even (including zero).
    """
    if type(chi) != str: raise TypeError("CHI must be a string")
    if len(chi) != 10: raise ValueError("CHI length must be 10")
    if int(chi[8])%2 == 0:
        return 'Female'
    else:
        return 'Male'
    ### END SOLUTION
```

```python
#artificially generated CHI numbers are provided here should you wish to use them to test your code
#You can call you function in this cell, but please only define it in the answer cell above

male_chi="0101011237"
female_chi="0101336489"

```


<img src="CHI_number_calc.png" alt="CHI_number" style="float: right; margin-right: 25px;" /> 

### CHI number validation

The digit in the 10th position is used to validate the first nine digits, to enable algorithmic checks to detect mistakes in manually entering the CHI number. 

Validation is achieved using the Modulus 11 algorithm, which is composed of four steps:

1. Starting from the left, multiply each of the first nine digits by (11 - digit position). For example, the value of the digit at  position 1 would be multiplied by 10, value of digit position 2 by 9, value of digit position 3 by 8, and so on.

2. Add the results of the nine multiplications together.

3. Divide the total by 11 and obtain the modulus (remainder).

4. Subtract the modulus (remainder) from 11 to get the calculated check digit. If the calculated check digit is 11, substitute a value of 0. If the calculated check digit is 10, then the first nine digits are invalid.


# Task 4

Write a Python 3 function called **`check_chi()`** that takes a 10 digit CHI number (as a string) as an argument and evaluates whether the CHI number is valid, according to the criteria listed below. 

If any of the following criteria are met, your function should return the string 'Invalid', otherwise, it should return 'Valid'. n.b. these strings are case sensitive.
* CHI number is not a string (1 mark)
* CHI number is <10 or >10 characters (1 mark)
* Check number is not valid based on the first nine digits. *Hint: you will need to calculate the expected check digit according to the information above* (4 marks)

Your function must have a valid docstring. (1 mark). Your function does **not** need to check that the 6 digit date of birth is a valid date. 

Your function must run without error to be awarded any marks for this question. To check the check digit, you will need to follow the instructions above to calculate the expected check digit, however you will score partial marks if your function successfully rejects invalid CHIs based on the other listed criteria but does not correct reject those with invalid check digits.

#### Mark scheme (7 marks total)
- Valid docstring - 1 mark
- Rejects CHIs with each of the listed rejection criteria - 1 or 4 marks per criterion as listed above, total of 6 marks available.


<!-- grade_id: Task4, points: 0, grade: false, solution: true -->
```python
def check_chi(chi): 
    ### BEGIN SOLUTION
    """
    Check CHI check digit is valid based length = 10, type=str, and modulus 11 check digit compatible 
    with digits 1-9 of the CHI number.
    """
    if type(chi) != str or len(chi)!=10: 
        return 'Invalid'
    position_weights_products=[(11-i)*int(x) for (i, x) in enumerate(chi[0:9], start=1)]
    sum_pw=sum(position_weights_products)
    check_digit=11-(sum_pw%11)
    if check_digit == 11: #if check digit is 11, set to 0
        check_digit = 0
    if int(chi[9])==check_digit:
        return 'Valid'
    else:
        return 'Invalid'
    ### END SOLUTION
```

```python
valid_chis=["0102674132", "1405769726", "2711087816", "0404763189", "2609661347", "0610826948"]

invalid_chis=["010156", 1405769726, "2711087811"]

```

# Autograders

Use the code cells below to test your code. Each autograder tests an individual aspect of the required performance of each function, therefore there are multiple autograders for each task.

Please note that there are no hidden tests - you will be awarded the marks for each of the autograders below that runs without error. If autograders do not pass, the error displayed is often not helpful in debugging your code: we recommend that you use the variables supplied within each task to test your code manually so that you can see the errors produced. 

### Task 1 autograders

<!-- grade_id: Task1_autograder1, points: 1, grade: true, solution: false -->
```python
#This grader tests that the function from task 1 returns unique species (1 mark)

species_list1=['G. propinqua', 'G. propinqua', 'G. propinqua', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis']

species_list2=['G. scandens', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. septentrionalis', 'G. septentrionalis', 'G. difficilis', 'G. acutirostris', 'G. acutirostris', 'G. difficilis', 'G. difficilis',  'G. difficilis',  'G. difficilis', 'G. scandens', 'G. scandens']

# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(sorted(unique_species(species_list1)), ['G. fortis', 'G. magnirostris', 'G. propinqua'])
assert_equal(sorted(unique_species(species_list2)), ['G. acutirostris', 'G. difficilis', 'G. fortis', 'G. fuliginosa', 'G. scandens', 'G. septentrionalis'])

print('This test has passed - list of unique species returned')
```

Output:
```
This test has passed - list of unique species returned
```

<!-- grade_id: Task1_autograder2, points: 1, grade: true, solution: false -->
```python
#This grader tests that the uniques species returns by the function from task 1 are sorted alphabetically (1 mark)

species_list1=['G. propinqua', 'G. propinqua', 'G. propinqua', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis']

species_list2=['G. scandens', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. septentrionalis', 'G. septentrionalis', 'G. difficilis', 'G. acutirostris', 'G. acutirostris', 'G. difficilis', 'G. difficilis',  'G. difficilis',  'G. difficilis', 'G. scandens', 'G. scandens']

# [Modify the tests below for your own problem]
# Check that squares returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(unique_species(species_list1), ['G. fortis', 'G. magnirostris', 'G. propinqua'])
assert_equal(unique_species(species_list2), ['G. acutirostris', 'G. difficilis', 'G. fortis', 'G. fuliginosa', 'G. scandens', 'G. septentrionalis'])

print('This test has passed - returned list is sorted alphabetically')
```

Output:
```
This test has passed - returned list is sorted alphabetically
```

<!-- grade_id: Task1_autograder3, points: 1, grade: true, solution: false -->
```python
# This autograder tests whether the function for task1 has a valid docstring
from nose.tools import assert_equal, assert_true
assert_true(unique_species.__doc__ is not None)
print('This test has passed - function has a docstring')
```

Output:
```
This test has passed - function has a docstring
```

### Task 2 autograders

<!-- grade_id: Task2_autograder1, points: 3, grade: true, solution: false -->
```python
#This autograder tests whether generate_dict_max() function from task 2 generated the expected dictionary
species_list1=['G. propinqua', 'G. propinqua', 'G. propinqua', 'G. propinqua', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. magnirostris', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis']
bill_length_list1=[19.5, 20.0, 22.5, 21.0, 25.0, 22.5, 21.5, 22.0, 24.0, 22.2, 21.0, 21.0, 22.5, 24.0, 22.0, 22.5, 24.0, 20.2, 20.0, 21.5, 17.0, 15.2, 16.0, 16.8, 16.0, 16.0, 16.5, 17.0, 16.5, 15.8, 15.8, 16.0, 18.5, 19.5, 15.8, 19.5, 17.0, 16.2, 16.0, 15.0, 15.0, 15.0, 16.5, 15.8, 13.0, 15.0, 15.0, 17.0, 15.2, 16.5]

species_list2=['G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fortis', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. fuliginosa', 'G. septentrionalis', 'G. septentrionalis', 'G. difficilis', 'G. acutirostris', 'G. difficilis',
 'G. difficilis', 'G. scandens', 'G. scandens', 'G. scandens', 'G. scandens', 'G. scandens', 'G. scandens']
tarsus_length_list2=[21.2, 23.0, 22.2, 22.0, 22.0, 23.0, 19.0, 18.2, 18.5, 20.0, 18.5, 20.5, 20.5, 19.0, 19.0, 19.0, 20.0, 20.2,
 22.0, 22.0, 20.5, 18.5, 24.0, 22.5, 21.5, 22.2, 20.2, 23.0, 22.0, 23.2]

from nose.tools import assert_equal
dict1=generate_dict_max(species_list1, bill_length_list1)
dict2=generate_dict_max(species_list2, tarsus_length_list2)
assert_equal(dict1['G. fortis'], 19.5)
assert_equal(dict1['G. propinqua'], 22.5)
assert_equal(dict1['G. magnirostris'], 25.0)
assert_equal(dict2['G. fuliginosa'], 20.5)
assert_equal(dict2['G. acutirostris'], 18.5)
assert_equal(dict2['G. septentrionalis'], 22.0)

print('This test has passed - Function generates a dictionary of maximum bill lengths (values) for each species (keys)')
```

Output:
```
This test has passed - Function generates a dictionary of maximum bill lengths (values) for each species (keys)
```

<!-- grade_id: Task2_autograder2, points: 1, grade: true, solution: false -->
```python
# This autograder tests whether the function for task2 has 6 lines of code or fewer (excluding docstring)
import inspect
from nose.tools import assert_equal, assert_true

#only award these marks if the function works
dict1=generate_dict_max(species_list1, bill_length_list1)
dict2=generate_dict_max(species_list2, tarsus_length_list2)
assert_equal(dict1['G. fortis'], 19.5)
assert_equal(dict1['G. propinqua'], 22.5)
assert_equal(dict1['G. magnirostris'], 25.0)
assert_equal(dict2['G. fuliginosa'], 20.5)
assert_equal(dict2['G. acutirostris'], 18.5)
assert_equal(dict2['G. septentrionalis'], 22.0)

lines = inspect.getsource(generate_dict_max).replace(';', '\n')
doc = inspect.getdoc(generate_dict_max).replace(';', '\n')
lines_filt = [line for line in lines.split('\n') if not line.lstrip().startswith(('#','"',"'")) and not len(line.lstrip())<1]
assert_true(len(lines_filt)<7)
print(f'This test has passed: 6 lines of fewer - {len(lines_filt)} lines counted')
```

Output:
```
This test has passed: 6 lines of fewer - 2 lines counted
```

<!-- grade_id: Task2_autograder3, points: 1, grade: true, solution: false -->
```python
# This autograder tests whether the function for task2 has 2 lines of code or fewer (excluding docstring)
import inspect
from nose.tools import assert_equal, assert_true

#only award these marks if the function works
dict1=generate_dict_max(species_list1, bill_length_list1)
dict2=generate_dict_max(species_list2, tarsus_length_list2)
assert_equal(dict1['G. fortis'], 19.5)
assert_equal(dict1['G. propinqua'], 22.5)
assert_equal(dict1['G. magnirostris'], 25.0)
assert_equal(dict2['G. fuliginosa'], 20.5)
assert_equal(dict2['G. acutirostris'], 18.5)
assert_equal(dict2['G. septentrionalis'], 22.0)

lines = inspect.getsource(generate_dict_max).replace(';', '\n')
doc = inspect.getdoc(generate_dict_max).replace(';', '\n')
lines_filt = [line for line in lines.split('\n') if not line.lstrip().startswith(('#','"',"'")) and not len(line.lstrip())<1]
assert_true(len(lines_filt)<3)
print(f'This test has passed: 2 lines of fewer: {len(lines_filt)} lines counted')
```

Output:
```
This test has passed: 2 lines of fewer: 2 lines counted
```

### Task 3 autograders

<!-- grade_id: Task3autograder10, points: 2, grade: true, solution: false -->
```python
# Tests whether function defined for task 3 correctly reads the 'Male' and 'Female' digits from valid CHI numbers
# Check that the function returns the correct output for several inputs:
from nose.tools import assert_equal
assert_equal(sex_from_chi("0102674132"), 'Male')
assert_equal(sex_from_chi("1405769726"), 'Female')
assert_equal(sex_from_chi("2711087816"), 'Male')
assert_equal(sex_from_chi("0404763189"), 'Female')
print('This test has passed: the function sex_from_chi() returned the correct output for different CHIs')
```

Output:
```
This test has passed: the function sex_from_chi() returned the correct output for different CHIs
```

<!-- grade_id: Task3autograder2, points: 2, grade: true, solution: false -->
```python
# Check that the function from Task 3 raises ValueErrors and TypeErrors error for invalid input:
from nose.tools import assert_raises
assert_raises(ValueError, sex_from_chi, "010267413")
assert_raises(ValueError, sex_from_chi, "01026741324")
assert_raises(TypeError, sex_from_chi, 3002649899)

print('This test has passed: the function sex_from_chi() produced ValueErrors or TypeError on incorrect input as requested in the question')
```

Output:
```
This test has passed: the function sex_from_chi() produced ValueErrors or TypeError on incorrect input as requested in the question
```

<!-- grade_id: Task3autograder3, points: 1, grade: true, solution: false -->
```python
# This autograder tests whether the function for task3 has a valid docstring
from nose.tools import assert_equal, assert_true
assert_true(sex_from_chi.__doc__ is not None)
print('This test has passed')
```

Output:
```
This test has passed
```

### Task 4 autograders

<!-- grade_id: Task4autograder1, points: 1, grade: true, solution: false -->
```python
# Task 4: Function checks CHI number is not a string (1 mark)
from nose.tools import assert_equal
assert_equal(check_chi("0102674132"), 'Valid')
assert_equal(check_chi(1405769726), 'Invalid')
assert_equal(check_chi(["0102674132"]), 'Invalid')
assert_equal(check_chi({"0102674132"}), 'Invalid')

print('This test has passed')
```

Output:
```
This test has passed
```

<!-- grade_id: Task4autograder2, points: 1, grade: true, solution: false -->
```python
# Task 4: Function checks CHI number is <10 or >10 characters (1 mark)
from nose.tools import assert_equal
assert_equal(check_chi("0102674132"), 'Valid')
assert_equal(check_chi("27110878165"), 'Invalid')
assert_equal(check_chi("271108"), 'Invalid')

print('This test has passed')
```

Output:
```
This test has passed
```

<!-- grade_id: Task4autograder3, points: 4, grade: true, solution: false -->
```python
# Task 4: Function checks check number is not valid based on the first nine digits. (4 marks)
from nose.tools import assert_equal
#Test valid CHI numbers
assert_equal(check_chi("0102674132"), 'Valid')
assert_equal(check_chi("1405769726"), 'Valid')
assert_equal(check_chi("2711087816"), 'Valid')
assert_equal(check_chi("2609661347"), 'Valid')
assert_equal(check_chi("0610826948"), 'Valid')
#Test CHI numbers with incorrect check digits
assert_equal(check_chi("0102674135"), 'Invalid')
assert_equal(check_chi("1405769721"), 'Invalid')
assert_equal(check_chi("2711087817"), 'Invalid')
assert_equal(check_chi("2609661343"), 'Invalid')
assert_equal(check_chi("0610826942"), 'Invalid')
print('This test has passed')
```

Output:
```
This test has passed
```

<!-- grade_id: Task4autograder4, points: 1, grade: true, solution: false -->
```python
# This autograder tests whether the function for task4 has a valid docstring
from nose.tools import assert_equal, assert_true
assert_true(check_chi.__doc__ is not None)
print('This test has passed')
```

Output:
```
This test has passed
```

```python

```